# Devoir maison 3 &ndash; Fonctions

## Rappels

### Deux façons de parcourir une liste

Nous avons vu deux façons de parcourir une liste : avec ou sans utiliser les indices.

```python
courses = ["biscuits", "fruits", "légumes", "viande", "bonbons"]
```

```python
# parcourir la liste sans utiliser les indices
for element in courses:
    print(element)
```

```python
# parcourir la liste en utilisant les indices
for i in range(len(courses)):
    print(courses[i])
```

### Deux façons de parcourir une chaîne de caractères

De même, nous pouvons parcourir une chaîne de caractères avec ou sans utiliser les indices.

```python
mot = "vacances"
```

```python
# parcourir une chaîne de caractères sans utiliser les indices
for lettre in mot:
    print(lettre)
```

```python
# parcourir une chaîne de caractères en utilisant les indices
for i in range(len(mot)):
    print(mot[i])
```

## Exercice 1 &ndash; Maximum

Écrire une fonction `maximum` :

- prenant en paramètre une liste **non vide** de nombres : `nombres` ;
- renvoyant le plus grand élément de cette liste.

Chacun des nombres utilisés est de type `int` ou `float`.

> On interdit ici d'utiliser `max`, ainsi que `sort` ou `sorted`.

<!-- grade_id: cell-a3da0d9c00ccba6c, points: 0, grade: false, solution: true -->
```python
def maximum(nombres):
    ### BEGIN SOLUTION
    M = nombres[0]
    for n in nombres:
        if n > M:
            M = n
    return M
    ### END SOLUTION
```

<!-- grade_id: cell-93244e687006c2c4, points: 2, grade: true, solution: false -->
```python
assert maximum([98, 12, 104, 23, 131, 9]) == 131
assert maximum([-27, 24, -3, 15]) == 24
### BEGIN HIDDEN TESTS
assert maximum([1, 3, -1, 4, 7, 0, 2]) == 7
assert maximum([10, 9, 1, -10, 12, 3, 1]) ==  12
### END HIDDEN TESTS
```

## Exercice 2 &ndash; Indice du minimum

Écrire une fonction `indice_min` qui prend en paramètre une liste **non vide** de nombres et qui renvoie l'indice de la première occurrence du minimum de cette liste.

> On n'utilisera pas les fonctions `min` et `index`.

<!-- grade_id: cell-01b0573239fe1c42, points: 0, grade: false, solution: true -->
```python
def indice_min(nombres):
    ### BEGIN SOLUTION
    m = nombres[0]
    i_min = 0
    for i in range(len(nombres)):
        if nombres[i] < m:
            m = nombres[i]
            i_min = i
    return i_min
    ### END SOLUTION
```

<!-- grade_id: cell-fad8a1ed5dac37ac, points: 2, grade: true, solution: false -->
```python
assert indice_min([5]) == 0
assert indice_min([2, 4, 1, 1]) == 2
assert indice_min([5, 3, 2, 5, 2]) == 2
### BEGIN HIDDEN TESTS
assert indice_min([3, -1, 2, -1, 4]) == 1
assert indice_min([4, 3, 3, 3, 3, 2, 3, 2]) == 5
### END HIDDEN TESTS
```

## Exercice 3 &ndash; Occurrences

Écrire une fonction `compte_occurrences` qui prend en paramètres `cible`, un caractère (une chaine de caractères de longueur 1), et `mot`, une chaine de caractères, et qui renvoie le nombre d'occurrences de `cible` dans `mot` ; c'est-à-dire le nombre de fois où `cible` apparait dans `mot`.

> On n'utilisera pas la méthode `count`.

<!-- grade_id: cell-78bdff27f3a77680, points: 0, grade: false, solution: true -->
```python
def compte_occurrences(caractere, mot):
    ### BEGIN SOLUTION
    n = 0
    for lettre in mot:
        if lettre == caractere:
            n += 1
    return n
    ### END SOLUTION
```

<!-- grade_id: cell-a22179ae26036156, points: 2, grade: true, solution: false -->
```python
assert compte_occurrences("o", "bonjour") == 2
assert compte_occurrences("a", "abracadabra") == 5
assert compte_occurrences("i", "abracadabra") == 0
### BEGIN HIDDEN TESTS
assert compte_occurrences("e", "helene") == 3
assert compte_occurrences("y", "kayak") == 1
### END HIDDEN TESTS
```